# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Pre-Lab #7**<span style="color:red"> $\to$ (4 possible marks)</span> </span></div> 

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Earth's Magnetic Field

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);"> Explore the duality between magnetism and electricity as described by Ampère's law and Faraday's law.  </span></b>
* <b><span style="color:rgba(0, 153, 76, 1);"> Develop an experimental method to measure the properties of Earth's magnetic field at our posiition on Earth's surface. </span></b>

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/images/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Pre-Lab 7.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
from IPython.display import YouTubeVideo
import math
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Introduction</span></div>

## Earth's magnetic field...

You are likely aware of the idea that the Earth has north and south magnetic poles. However, these names are a bit misleading. For one thing, since Earth's geographic north attracts the north pole of a bar magnet, Earth's *geographic* north pole is  a *magnetic* south pole.  This situation is depicted in **Fig. 1**.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/Earth-field.jpg" alt="Earth's magnetic field" style="width: 45%; height: auto;" /><br>
    <b>Fig. 1: The <i>magnetic</i> south pole of the Earth is located at its <i>geographic</i> north pole.    Figure adapted from <a href ="https://commons.wikimedia.org/wiki/File:Openstax_college-physics_22.4_earth-magnet.jpg">Wikimedia Commons</a>.</b>
</center>
</p>

Further, these poles don't exactly coincide with the geological poles, and in fact the north pole is wandering NNW at about 55km per year. To confuse matters even more, by taking note of how flows of ferromagnetic magma have solidified in the past, we know that the Earth's magnetic field undergos occasional **geomagnetic reversals**, in which the magnetic poles effectively swap places.  

Aside from all these complications, the magnetic field of the Earth is stable enough for us to use it for navigation. In a compass, a magnetized needle on a low-friction bearing will swivel to align itself with the local magnetic field. This tells us what *direction* the horizontal component of the local field is pointing in, but what of the *magnitude*? Where you stand, how strong is the Earth's magnetic field?

In this pre-lab, you will use simple vector analysis to see how it's possible to measure the local magnetic field using only a compass, a coil of wire, and a DC power source. This will equip you with the analytic tools to carry out the actual experiment in this week's lab.  

## Creating uniform magnetic fields...

The technique that we'll use to probe the magnetic field of the Earth requires use to generate a a uniform magnetic field and control its strength.  Before considering how we can produce a uniform magnetic field, let's first recall that capacitors can be used to produce uniform electric fields.  The left-had side of **Fig. 2** shows the electric field produced by a charged parallel-plate capacitor.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/uniformE.png" alt="Uniform electric field" style="width: 100%; height: auto;" /><br>
    <b>Fig. 2: The electric field between the plates of a charged capacitor is approximately uniform.  Figure adapted from <a href ="https://commons.wikimedia.org/wiki/File:VFPt_capacitor-round-plate.svg">Wikimedia Commons</a>.</b>
</center>
</p>

Notice that:
 - The electric field is strong between the plates and relatively weak everywhere else
 - The electric field between the plates is approximately uniform, especially in the region near the centre of the capacitor
 
It is common to approximate the electric field of a capacitor as shown onm the right-hand side of **Fig. 2**.  In this approximation, the electric field between the plates is uniform with a magnitude given by $E=\sigma/\varepsilon_0$, where $\sigma$ is the charge per unit area on the plates of the capacitor.  The electric field everywhere else is taken to be zero.  This approximation is very good when the plates are large compared the gap between the two plates. 
 
Whereas charge is a source of electric fields, we will see in class that *current* (or charge in motion) is a source of magnetic fields.  We can generate a uniform magnetic field using a coil of wire carrying a constant current $I$.  The left-hand side of **Fig. 3** shows the magnetic field created by a cylindrical coil supplied with a constant current.  These coils are often called **solenoids** or **inductors**.  

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/uniformB.png" alt="Uniform magnetic field" style="width: 100%; height: auto;" /><br>
    <b>Fig. 3: The magnetic field inside a solenoid carrying current $I$ is approximately uniform.  Figure adapted from <a href ="https://commons.wikimedia.org/wiki/File:VFPt_magnet_B_currents.svg">Wikimedia Commons</a>.</b>
</center>
</p>
 
Notice that:
 - The magnetic field is strong within the bore of the solenoid and relatively weak everywhere else
 - The magnetic field between the plates is approximately uniform, especially in the region near the centre of the solenoid
 
Like we did for the electric field of the capacitor, it is common to approximate the magnetic field of a solenoid as shown on the right-hand side of **Fig. 3**.  In this approximation, the magnetic field in the bore of the solenoid is uniform and its zero outside the bore.  After studying Ampère's law, we'll be able to show that the magnitude of the uniform magnetic field in the bore is given by:

\begin{align}
B_\mathrm{S} &= \frac{\mu_0 N I}{\ell}\\
&= \mu_0 n I\tag{1}
\end{align}

In Eq. (1):

 - $\mu_0 = 4\pi\times 10^{-7}\rm\ N/A^2$ is a constant called the *permeability of free space*
 - $N$ is the number of windings or loops of wire used to form the solenoid and $\ell$ is the length of the solenoid
 - $n = N/\ell$ is the density of windings or the number of loops per unit length
 
The approximation shown on the right-hand side of **Fig. 3** is very good when then solenoid's length is very long compared to its radius.  When this condition is satisfied, Eq. (1) is valid for points that are not too close to the ends of the solenoid.

It is important to emphasize that solenoids are used all of time in practical applications that require uniform magnetic fields.  In these applications, the coils of current-carrying wires are often called *electromagnets*.  Magnetic Resonance Imaging (MRI) scanners use large electromagnets wound with superconducting wires that are cooled to $4\rm\ K$ ($-269^\circ\rm C$) using liquid helium.  The bore of the electromagnet is large enough for a person to lie inside of it.  **Figure 4** shows a photograph of one of the first MRI scanners.  This electromagnet of this scanner wasn't cooled to low temperature or wound using superconducting wires.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/MRI.png" alt="MRI" style="width: 40%; height: auto;" /><br>
    <b>Fig. 4: An early MRI scanner showing the exposed electromagnet/solenoid used to produce a large and approximately uniform magnetic field.  Figure adapted from <a href ="https://commons.wikimedia.org/wiki/File:MRI_Scanner_Mark_One.jpg">Wikimedia Commons</a>.</b>
</center>
</p>

As a second example, small solenoids are used in the design of stepper motors which very good control over the motor's speed and position.  See **Fig. 5**.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/stepper.png" alt="Stepper motor" style="width: 40%; height: auto;" /><br>
    <b>Fig. 5: (Left) The conceptual design of a stepper motor.   (Right) A partially disassembled stepper motor with some of the solenoid winding exposed.  Figures taken from <a href ="https://freesvg.org/schrittmotor-01-c">FreeSVG</a> and <a href ="https://commons.wikimedia.org/wiki/File:Stepper_motor_1.jpg">Wikimedia Commons</a>.</b>
</center>
</p>

 
## Creating uniform magnetic fields...
 
In Lab #7, you will place a compass inside a solenoid.  If the current in the solenoid is zero, then the compass needle will align with the direction of the horizontal component of Earth's magnetic field.  This situation is depicted in **Fig. 6** which shows a compass placed in the centre of the bore of a solenoid.  Because the current is zero, the orientation of the compass needle is determined entirely by Earth's magnetic field $\vec{B}_\mathrm{E}$

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/solenoid_Izero.jpg" alt="Compass in a solenoid" style="width: 40%; height: auto;" /><br>
    <b>Fig. 6: A compass inside a solenoid with $I=0$.  Figure adapted from images taken from <a href ="https://freesvg.org/compass-simple">FreeSVG</a> and <a href ="https://commons.wikimedia.org/wiki/File:VFPt_magnet_B_currents.svg">Wikimedia Commons</a>.</b>
</center>
</p>

Next, a small current is supplied to the solenoid.  In this case, the solenoid produces a uniform magnetic field $\vec{B}_\mathrm{S}$ that is perpendicular to $\vec{B}_\mathrm{E}$.  The compass needle now aligns with the net magnetic field:

$$
\vec{B}_\mathrm{net} = \vec{B}_\mathrm{E} + \vec{B}_\mathrm{S}
$$

which makes an angle $\varphi$ with respect to $\vec{B}_\mathrm{E}$.  See **Fig. 7**.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/solenoid_Iweak.jpg" alt="Compass in a solenoid" style="width: 40%; height: auto;" /><br>
    <b>Fig. 7: A compass inside a solenoid with a small current $I$.  Figure adapted from images taken from <a href ="https://freesvg.org/compass-simple">FreeSVG</a> and <a href ="https://commons.wikimedia.org/wiki/File:VFPt_magnet_B_currents.svg">Wikimedia Commons</a>.</b>
</center>
</p>

Of course, as depicted in **Fig. 8**, increasing the current $I$ increases the strength of $\vec{B}_\mathrm{S}$ which causes the compass needle to rotate further away from the direction of $\vec{B}_\mathrm{E}$.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Pre-Lab-7/images/solenoid_Istrong.jpg" alt="Compass in a solenoid" style="width: 40%; height: auto;" /><br>
    <b>Fig. 8: A compass inside a solenoid with a large current $I$.  Figure adapted from images taken from <a href ="https://freesvg.org/compass-simple">FreeSVG</a> and <a href ="https://commons.wikimedia.org/wiki/File:VFPt_magnet_B_currents.svg">Wikimedia Commons</a>.</b>
</center>
</p>

In Lab #7, you will measure the angle $\varphi$ that the compass needle makes with respect to $\vec{B}_\mathrm{E}$ as a function of the current $I$ supplied to the solenoid.  You will use this data to determine the strength of the horizontal component of Earth's magnetic field.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Magnetic Field Calculation</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>**  **<span style="color:red">(1 mark)</span>**

Suppose you place a compass in the centre of a solenoid as shown in **Fig. 7**.  The axis of the solenoid is perpendicular to Earth's magnetic field and carries a nonzero current.  The compass needle makes an angle $\varphi$ with respect to $\vec{B}_\mathrm{E}$.  Suppose now that the compass is pulled away from the centre of the solenoind and placed close to one of the ends.  Based on the magnetic field shown on the left-hand side of **Fig. 3**, would the angle $\varphi$ increase or decrease?  Explain your reasoning. 

***
**<span style="color:blue">Answer 2.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 
***

<!-- END QUESTION -->

Below, we define a set of variable which will be used to solve **Question 2.2**. `Be` represents the magnitude of Earth's magnetic field $B_\mathrm{E}$, `Bs` is the magnitude of the solenoid's magnetic field $B_\mathrm{S}$, and `phi` is the angle $\varphi$ that the compass needle makes with $\vec{B}_\mathrm{E}$. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>**  **<span style="color:red">(1 mark)</span>**

The net magnetic field at the position of the compass is:

$$ \vec{B}_\mathrm{net} = \vec{B}_\mathrm{E} + \vec{B}_\mathrm{S} $$

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/6f1d0e7429284a6d30eb0dfff436873fbe37228b/Pre-Lab-7/images/btot.svg" style="width: 15%" /><br>
<b>Fig. 9: The net magnetic field is given by superposition of Earth's magnetic field $\vec{B}_\mathrm{E}$ and the solenoid's magnetic field $\vec{B}_\mathrm{S}$. The direction of the net field can be changed by varying the current in the solenoid. </b>
</center>
</p>


The compass needle will point along the direction of $\vec{B}_\mathrm{net}$. Noting that $\vec{B}_\mathrm{E}\perp\vec{B}_\mathrm{S}$, use some simple trigonometric analysis to find an expression relating the angle $\varphi$ in **Fig. 9** above and the magnitudes of the two magnetic fields: $B_\mathrm{E} = \left\vert \vec{B}_\mathrm{E}\right\vert$ and $B_\mathrm{S} = \left\vert \vec{B}_\mathrm{S}\right\vert$. Solve your expression for $B_\mathrm{E}$ and enter the result below.

***
**<span style="color:blue">Answer 2.2:</span>**

Using the cell below, type your expression for $B_\mathrm{E}$ (denoted `Be` below) in terms of the variables `Bs`, and `phi`.

To express your answer, we are using a package called ```sympy```. This allows us to symbolically represent a variable, rather than having them assigned to numerical values. We are also importing ```tan``` from this package to allow for symbolic representation of this trigonometric function. You can type ```tan(phi)```, to symbolically represent $\tan(\varphi)$ in the code below.

```python
# Import the sympy package
from sympy import symbols, tan

# Please do not delete or modify the line directly below this comment.  
# It is used to allow the user to enter symbol expressions into this cell.
Be, Bs, phi = symbols('Be Bs phi')

# Enter a symbolic expression for B_loc in terms of B_sol and phi.
Be = ...
Be
```

```python
grader.check("q2.2")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>**  **<span style="color:red">(1 mark)</span>**

If we're thinking like experimentalists, we should be seeking an expression of the quantity of interest (Earth's magnetic field) in terms of easily measured quantities. The angle $\varphi$ is easy to measure, but it's not obvious how to measure $B_\mathrm{S}$. 

Recall Eq. (1) for the magnetic field of solenoid:

$$
B_\mathrm{S} = \mu_0 n I
$$

Use this result and your expression from **Question 2.2** to find an expression for $B_\mathrm{E}$ in terms of the solenoid current $I$.  You expression should also involve the constant $\mu_0$ and the density of turns $n$.

***
**<span style="color:blue">Answer 2.3:</span>**

In the cell below, type a symbolic expression for `Be` in terms of `phi`, `I`, `n`, and `u0`. Remember that $\mu_0$ just represents a constant equal to $4\pi \times 10^{-7}\rm\ N/A^2$.

```python
# Please do not delete or modify the line directly below this comment.  
# It is used to allow the user to enter symbol expressions into this cell.
Be, n, I, u0 = symbols('Be n I u0')

# Enter a symbolic expression for Be in terms of I, n, u0, and phi.
Be = ...
Be
```

```python
grader.check("q2.3")
```

<!-- END QUESTION -->

***
**<span style="color:blue">Question 2.4</span>**  **<span style="color:red">(1 mark)</span>**

To get a better measurement of $B_\mathrm{e}$, we should take a set of measurements of $\varphi$ over a range of currents $I$. The expression from **Question 2.3** allows us to express $\tan(\varphi)$ as a function of the solenoid current $I$:

$$
\tan(\varphi) = \left(\frac{\mu_0 n}{B_\mathrm{E}}\right)I
\tag{2}
$$


Suppose you were to plot $\tan\varphi$ as a function of $I$ -- the result would be a linear set of data with a $y$-intercept of zero as shown in **Fig. 10**. 

<br>
<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab7/images/linearPlot.png" style="width: 35%" />
<br>
<b>Fig. 10: An example plot of $\tan\varphi$ as a function of $I$ (using made up data).  The red line represents a linear fit to the data points. </b></center>
</p>

The slope $m$ of the $\tan(\varphi)$ versus $I$ plot is:

$\begin{align}
\mathrm{(a)~}  m &= \dfrac{\mu_0 n}{B_\mathrm{E}}\\
\mathrm{(b)~}  m &= -\dfrac{\mu_0 n}{B_\mathrm{E}}\\
\mathrm{(c)~}  m &= \dfrac{1}{B_\mathrm{E}}\\
\mathrm{(d)~}  m &= -\dfrac{1}{B_\mathrm{E}}\\
\mathrm{(e)~}  m &= \dfrac{B_\mathrm{E}}{\mu_0 n}\\
\mathrm{(f)~}  m &= -\dfrac{B_\mathrm{E}}{\mu_0 n}
\end{align}$

***
**<span style="color:blue">Answer 2.4:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, e, or f) between single or double quotes.  

*** Please do not change anything to the left of the equals sign. ***

```python
a2_4 = ...
```

```python
grader.check("q2.4")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 121 pre-labs!  Please feel free to include any comments you have about this pre-lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs and pre-labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the pre-lab, ...

This question is optional and will have no impact on your pre-lab or lab grade.

***
**<span style="color:blue">Answer 3.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Lab Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  For example, you could compare Gaussian distributions with different standard deviations and/or means. 

```python
# Here's an empty code cell that you can use.

```

```python
# Here's another empty code cell that you can use.

```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.

```

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: March 13, 2023

Tests for q2.2 (points: None)
```
>>> Be == Bs/tan(phi)
True
```
failure_message: Note that Be and Bs form a right-angle triangle.  Be is the side adjacenet to the angle phi and Bs is opposite phi.


Tests for q2.3 (points: None)
```
>>> Be == u0*n*I/tan(phi)
True
```
failure_message: Substitute Bs = u0*n*I into your solution to Question 2.2


Tests for q2.4 (points: None)
```
>>> a2_4.lower() in ["a", "b", "c", "d", "e", "f"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', 'e' or 'f'.
```
>>> PHYS121.get_hash(a2_4.lower()) == '0cc175b9c0f1b6a831c399e269772661'
True
```
failure_message: Equation (2) is of the form tan(phi) = m*I, where m is the slope.
